# Структуры данных

Работая с большим количеством одинаковых данных их удобно хранить в структурированном виде, что позволяет обращатся к ним систематизированно<br>

**Julia** поддерживает следующие структуры данных:
1. Кортежи (Tuples)
2. Словари (Dictionaries)
3. Массивы (Arrays)

<br>
Коротко их можно охарактеризовать следующим образом - массивы и кортежи являются упорядаченными структурами (доступ к элементам можно получить по их номеру), словари и массивы - редактируемые структуры.

## Кортежи

Мы можем создать кортеж путем помещения упорядоченой коллекции элементов в круглые скобки `( )`.

Синтаксис: <br>
```julia
(элемент1, элемент2, ...)
```


```julia
myfavoriteanimals = ("Пингвины", "Коты", "Сахарные летяги")
```

Output:
```
("Пингвины", "Коты", "Сахарные летяги")
```

Доступ к элемента осуществляется по их индексу,

```julia
myfavoriteanimals[1]
```

Output:
```
"Пингвины"
```

но поскольку кортежи нередактируемые, мы не можем изменить их элементы

```julia
myfavoriteanimals[1] = "Вомбаты"
```

Output:
```
Stacktrace:

 [1] top-level scope

   @ c:\Users\edaev\Desktop\MIOiS_lessons-main\MIOiS_lessons-main\Lesson_1\jl_notebook_cell_df34fa98e69747e1a8f8a730347b8e2f_W6sZmlsZQ==.jl:1
```

##  Именнованные кортежи

Начиная с версии 1.6 **Julia** поддерживает именнованные массивы, в которых каждый элемент кроме номера также имеет свое имя. Для добавления имени необходимо использовать `=` внутри кортежа:

```julia
(name1 = item1, name2 = item2, ...)
```

```julia
myfavoriteanimals_1 = (bird = "Пингвины", mammal = "Коты", marsupial = "Сахарные летяги")
```

Output:
```
(bird = "Пингвины", mammal = "Коты", marsupial = "Сахарные летяги")
```

Как и обычные кортежи, именнованные кортежи позволяют обращатся к элементам по номеру:

```julia
myfavoriteanimals_1[1]
```

Output:
```
"Пингвины"
```

Но так же к ним можно образатся по имени:

```julia
myfavoriteanimals_1.bird
```

Output:
```
"Пингвины"
```

## Словари

Для хранения набора парных данных удобно использовать словари. Для создания словаря необходимо использовать функцию `Dict()`.

Синтаксис:
```julia
Dict(ключ1 => значение1, ключ2 => значение2, ...)
```

Хороший пример - записная книжка, где каждому имени соответствует номер телефона.

```julia
myphonebook = Dict("Деннис" => "+7(977) 53-09-668", "Охотники за превидениями" => "555-2368")
```

Output:
```
Dict{String, String} with 2 entries:
  "Деннис"                   => "+7(977) 53-09-668"
  "Охотники за превидениями" => "555-2368"
```

В данном примере каждые имя и номер являются парой "ключ" и "значение". Мы можем узнать номер Денниса используяю соответствующий ключ

```julia
myphonebook["Деннис"]
```

Output:
```
"+7(977) 53-09-668"
```

Мы можем добавить новую запись в словарь просто путем присваивания

```julia
myphonebook["Жди меня"] = "8-800-700-84-36"
```

Output:
```
"8-800-700-84-36"
```

Теперь записная книжка выглядит следующим образом

```julia
myphonebook
```

Output:
```
Dict{String, String} with 3 entries:
  "Жди меня"                 => "8-800-700-84-36"
  "Деннис"                   => "+7(977) 53-09-668"
  "Охотники за превидениями" => "555-2368"
```

Удаление элементов осуществляется при помощи функции `pop!`

```julia
pop!(myphonebook, "Жди меня")
```

Output:
```
"8-800-700-84-36"
```

```julia
myphonebook
```

Output:
```
Dict{String, String} with 2 entries:
  "Деннис"                   => "+7(977) 53-09-668"
  "Охотники за превидениями" => "555-2368"
```

В отличии от кортежей и массивов, словари не упорядоченны, поэтому мы не можем обращатся к их элементам по индексу

```julia
myphonebook[1]
```

Output:
```
@ Base .\dict.jl:498

 [2] top-level scope

   @ c:\Users\edaev\Desktop\MIOiS_lessons-main\MIOiS_lessons-main\Lesson_1\jl_notebook_cell_df34fa98e69747e1a8f8a730347b8e2f_X34sZmlsZQ==.jl:1
```

В примере выше **Julia** считает, что вы пытаетесь обратиться к элементу с ключем `1`.

Для получения списка всех доступных ключей можно использовать функцию ```keys```

```julia
keys(new_dict)
```

Output:
```
Stacktrace:

 [1] top-level scope

   @ c:\Users\edaev\Desktop\MIOiS_lessons-main\MIOiS_lessons-main\Lesson_1\jl_notebook_cell_df34fa98e69747e1a8f8a730347b8e2f_X40sZmlsZQ==.jl:1
```

## Массивы

Как уже говорилось, в отличии от кортежей, массивы могут быть изменены, а в отличии от словарей - данные в массиве упорядочены.<br>
Самым простым способо определить массив является перечисление его элементов в квадратных скобках `[ ]`.

Синтаксис: <br>
```julia
[элемент1, элемент2, ...]
```

К примеру, мы можем создать массив студентов

```julia
students = ["Иванов", "Петров", "Васечкин", "Сидоров", "Солтыков-Щедрин"]
```

Output:
```
5-element Vector{String}:
 "Иванов"
 "Петров"
 "Васечкин"
 "Сидоров"
 "Солтыков-Щедрин"
```

После того, как массив создан, мы моем получить доступ к его элемента по их индексу. Индексация в **Julia** начинается с 1.

```julia
students[3]
```

Output:
```
"Васечкин"
```

Таким же образом, мы можем менять его элементы

```julia
students[3] = "Соломоненко"
```

Output:
```
"Соломоненко"
```

```julia
students
```

Output:
```
5-element Vector{String}:
 "Иванов"
 "Петров"
 "Соломоненко"
 "Сидоров"
 "Солтыков-Щедрин"
```

Так же, имеется несколько функции определяющих массив заданной формы:

```julia
ones_ = ones(4)
```

Output:
```
4-element Vector{Float64}:
 1.0
 1.0
 1.0
 1.0
```

```julia
zeros_ = zeros(4,2)
```

Output:
```
4×2 Matrix{Float64}:
 0.0  0.0
 0.0  0.0
 0.0  0.0
 0.0  0.0
```

```julia
rands_ = rand(2,3,1)
```

Output:
```
2×3×1 Array{Float64, 3}:
[:, :, 1] =
 0.106653  0.767319  0.343381
 0.387152  0.893063  0.0535204
```

Эти функции работают схожим образом - они возвращают массив заданного размера, состоящий из единиц, нулей или случайных чисел.<br>
Как можно было заметить, тип у них отличается: в первом случае это `Vector{Float64}`, во втором `Matrix{Float64}` и в третьем - `Array{Float64,3}`. На самом деле `Vector{T}` и `Matrix{T}` - это всего лишь псевдонимы для `Array{T,1}` и `Array{T,2}`, где `T` - это тип компонентов массива.

В **Julia** полезно заранее объявлять массив - это позволяет оптимизировать вычисления. Для этого существует универсальный конструктор массивов 

```julia
arr_1 = Array{Int64}(undef, 3,5)
```

Output:
```
3×5 Matrix{Int64}:
  85899345940  3018925951920  3018954030384  3018948832912  281474976776449
 133143986202  3018948832656  3018954030480  3018948832976                1
 146028888096  3018954030288  3018948832848  3018954030240    3018966907968
```

Здесь `undef` - ключевое слово, которое означает, что массив элементы массива не иницализируются при ео объявлении. Сразу после объявления в них будет находиться "*мусор*" из памяти.<br>
В объявленный таким образом массив уже не получится занести элемент другой типа:

```julia
arr_1[1,1] = 3.5
```

Output:
```
@ Base .\array.jl:1024

 [4] top-level scope

   @ c:\Users\edaev\Desktop\MIOiS_lessons-main\MIOiS_lessons-main\Lesson_1\jl_notebook_cell_df34fa98e69747e1a8f8a730347b8e2f_X62sZmlsZQ==.jl:1
```

В случае, если заранее не известно, какого типа элементы могут храниться в массиве, в качестве типа можно указать `Any`

```julia
arr_2 = Array{Any}(undef, 4,4)
```

Output:
```
4×4 Matrix{Any}:
 #undef  #undef  #undef  #undef
 #undef  #undef  #undef  #undef
 #undef  #undef  #undef  #undef
 #undef  #undef  #undef  #undef
```

```julia
fibonacci = Array{Float64}([1, 1, 2, 3, 5, 8, 13])
```

Output:
```
7-element Vector{Float64}:
  1.0
  1.0
  2.0
  3.0
  5.0
  8.0
 13.0
```

При создании массива с содержимым разного типа, ему будет автоматически присвоен тип `Array{Any}`

```julia
mixture = [1, 1, 2, 3, "Ted", "Robyn"]
```

Output:
```
6-element Vector{Any}:
 1
 1
 2
 3
  "Ted"
  "Robyn"
```

Для редактирования массивов так же имеются функции `push!` и `pop!`. `push!` добавляет элементы в конец массива, а `pop!` удаляет последний элемент из массива.

```julia
push!(fibonacci, 21)
```

Output:
```
8-element Vector{Float64}:
  1.0
  1.0
  2.0
  3.0
  5.0
  8.0
 13.0
 21.0
```

```julia
pop!(fibonacci)
```

Output:
```
21.0
```

```julia
fibonacci
```

Output:
```
7-element Vector{Float64}:
  1.0
  1.0
  2.0
  3.0
  5.0
  8.0
 13.0
```

Поскольку массив - это тоже тип данных, мы можем определять массив массивов:

```julia
favorites = [["koobideh", "chocolate", "eggs"],["penguins", "cats", "sugargliders"]]
```

Output:
```
2-element Vector{Vector{String}}:
 ["koobideh", "chocolate", "eggs"]
 ["penguins", "cats", "sugargliders"]
```

При этом, они не обяаны быть совместимй размерности:

```julia
numbers = [[1, 2, 3], [4, 5], [6, 7, 8, 9]]

```

Output:
```
3-element Vector{Vector{Int64}}:
 [1, 2, 3]
 [4, 5]
 [6, 7, 8, 9]
```

Однако не следует путать их с двумерными массивами (или матрицами) - доступ к элементам осуществляется по разному

```julia
numbers[1][2]
```

Output:
```
2
```

```julia
numbers[1,2]
```

Output:
```
@ Base .\essentials.jl:14

 [2] top-level scope

   @ c:\Users\edaev\Desktop\MIOiS_lessons-main\MIOiS_lessons-main\Lesson_1\jl_notebook_cell_df34fa98e69747e1a8f8a730347b8e2f_Y114sZmlsZQ==.jl:1
```

```julia
zeros_[2,2]
```

Output:
```
0.0
```

```julia
zeros_[2][2]
```

Output:
```
@ Base .\number.jl:98

 [2] top-level scope

   @ c:\Users\edaev\Desktop\MIOiS_lessons-main\MIOiS_lessons-main\Lesson_1\jl_notebook_cell_df34fa98e69747e1a8f8a730347b8e2f_Y116sZmlsZQ==.jl:1
```

**Будте остороны при копировании массивов!**

```julia
fibonacci
```

Output:
```
7-element Vector{Float64}:
  1.0
  1.0
  2.0
  3.0
  5.0
  8.0
 13.0
```

```julia
somenumbers = fibonacci
```

Output:
```
7-element Vector{Float64}:
  1.0
  1.0
  2.0
  3.0
  5.0
  8.0
 13.0
```

```julia
somenumbers[1] = 404
```

Output:
```
404
```

```julia
fibonacci
```

Output:
```
7-element Vector{Float64}:
 404.0
   1.0
   2.0
   3.0
   5.0
   8.0
  13.0
```

Редактирование `somenumbers` привело к изменению переменной `fibonacci`!

В данном примере мы, в де1ствительности, не копировали переменную `fibonacci`. Мы просто создали новую ссылку для доступа к объекту `fibonacci`.

Если нам необходимо создать компию массива `fibonacci`, необходимо использовать функцию `copy`.

```julia
# Для начала, восcтановим fibonacci
fibonacci[1] = 1
fibonacci
```

Output:
```
7-element Vector{Float64}:
  1.0
  1.0
  2.0
  3.0
  5.0
  8.0
 13.0
```

```julia
somemorenumbers = copy(fibonacci)
```

Output:
```
7-element Vector{Float64}:
  1.0
  1.0
  2.0
  3.0
  5.0
  8.0
 13.0
```

```julia
somemorenumbers[1] = 404
```

Output:
```
404
```

```julia
fibonacci
```

Output:
```
7-element Vector{Float64}:
  1.0
  1.0
  2.0
  3.0
  5.0
  8.0
 13.0
```

### Упражнения

#### 3.1 
Создайте массив `a_ray` с элементами `[1, 2, 3]`

```julia
a_ray = [1, 2, 3]
```

Output:
```
3-element Vector{Int64}:
 1
 2
 3
```

Добавьте `4` в конец массива.

```julia
push!(a_ray, 4)
```

Output:
```
4-element Vector{Int64}:
 1
 2
 3
 4
```

```julia
@assert a_ray == [1, 2, 3, 4]
```

и удалите ее

```julia
pop!(a_ray)
```

Output:
```
4
```

<!-- grade_id: cell-a6b20b78fbfeb449, points: 1, grade: true, solution: false -->
```julia
@assert a_ray == [1, 2, 3]
```

#### 3.2 
Добавьте в записную книжку `myphonebook` телефон пожарной `101`?

```julia
myphonebook["пожарная"] = "101"
```

Output:
```
Dict("Деннис" => "+7(977) 53-09-668", "Охотники за превидениями" => "555-2368", "пожарная" => "101")
```

#### 3.3 
Создайте переменную `flexible_phonebook` , в которой могли бы хранится номера как в строковом, так и в численном виде, и добавьте туда номера *Денниса* и *Охотников за превидениями*.



```julia
flexible_phonebook = Dict{String, Any}("Dennis" => "8-800-555-35-35", "Ohotniki_za_PRIVIDENIYAMI" => "89999454556")
```

Output:
```
Dict{String, Any} with 2 entries:
  "Dennis"                    => "8-800-555-35-35"
  "Ohotniki_za_PRIVIDENIYAMI" => "89999454556"
```

```julia
@assert typeof(flexible_phonebook) == Dict{String, Any}
```