# Основы линейной алгебры в Julia|
Author: Andreas Noack Jensen (MIT & JuliaComputing) (https://twitter.com/anoackjensen?lang=en)

Для начала, давайте объявим матрицу случайных величин

```julia
A = rand(1:4,3,3)
```

Output:
```
3×3 Matrix{Int64}:
 3  1  4
 4  2  2
 1  3  4
```

Объявим вектор единиц

```julia
x = fill(1.0, (3,)) # = fill(1.0, 3)
```

Output:
```
3-element Vector{Float64}:
 1.0
 1.0
 1.0
```

Обратите внимание, что $A$ имеет тип `Matrix{Int64}`, в то время как $x$ имеет тип `Vector{Float64}`. В Julia определены псевдонимы для `Vector{Type}=Array{Type,1}` и `Matrix{Type}=Array{Type,2}`. 

Многие базовые операторы выглядят так же, как в других языках програмирования.
#### Умножение

```julia
b = A*x
```

Output:
```
3-element Vector{Float64}:
 8.0
 8.0
 8.0
```

#### Транспорование
как и в *Matlab* `A'` возвращает Эрмитово-сопряженную матрицу (матрицу сопряженно-транспонированную)

```julia
A'
```

Output:
```
3×3 adjoint(::Matrix{Int64}) with eltype Int64:
 3  4  1
 1  2  3
 4  2  4
```

```julia
A
length(A[1,1:end])
```

Output:
```
3
```

```julia
x_5 = 5
x_5 = x+1
complex(1,3)
```

Output:
```
Stacktrace:

 [1] top-level scope

   @ c:\Users\edaev\Desktop\MIOiS_lessons-main\Lesson_4\jl_notebook_cell_df34fa98e69747e1a8f8a730347b8e2f_X13sZmlsZQ==.jl:2
```

Обычное транспонирование осуществляется при помощи

```julia
transpose(A)
```

Output:
```
3×3 transpose(::Matrix{Int64}) with eltype Int64:
 3  4  1
 1  2  3
 4  2  4
```

#### Транспонированное умножение
*Julia* позволяет записать без *

```julia
A'A
```

Output:
```
3×3 Matrix{Int64}:
 26  14  24
 14  14  20
 24  20  36
```

#### решение СЛАУ
Уравнение $Ax=b$ для ***квадратной*** $A$ решается при помощи оператора \.

```julia
b' /A
```

Output:
```
1×3 adjoint(::Vector{Float64}) with eltype Float64:
 -0.5  2.0  1.5
```

```julia

```

Также `A\b` возвращает *оптимальное в среднеквадратичном смысле* если у нас переопределенная система (т.е. матрица `A` "высокая")

```julia
Atall = rand(3, 2)
```

Output:
```
3×2 Matrix{Float64}:
 0.588143  0.357542
 0.650556  0.779886
 0.166704  0.973618
```

```julia
Atall\b
```

Output:
```
2-element Vector{Float64}:
 7.173737393443174
 6.280218644041781
```

Если же система недоопределена, то оператор `\` возвращает  *решение с минимальной евклидовой нормой*

```julia
bshort = rand(2)
Ashort = rand(2, 3)
```

Output:
```
2×3 Matrix{Float64}:
 0.676813  0.328247  0.985726
 0.339542  0.946853  0.636701
```

```julia
Ashort\bshort
```

Output:
```
3-element Vector{Float64}:
 0.2850766332253017
 0.07553055112814047
 0.4037892937032488
```

# Модуль LinearAlgebra

Хоть и многие инструменты линейной алгебры доступны в Julia по умолчанию, существет стандартная библиотека `LinearAlgebra` которая предоставляет еще больше довольно важных функций. В частности, в ней имеются инструменты факторизации и некоторые типы для матриц определенного типа.

### Упражнения

#### 10.1 
Вычислите скалярное произведение вектора `v` с самиси собой и запишите результат в переменную `dot_v`.



```julia
using LinearAlgebra
v = [1,2,3]
dot_v = dot(v, v)
```

Output:
```
14
```

```julia

```

<!-- grade_id: cell-913fef9b0d19cd52, points: 1, grade: true, solution: false -->
```julia
@assert dot_v == 14
```

#### 10.2 
Вычислите тензорное произведение вектора с самим собой и запишите полученную матрицу в переменную `outer_v`
$$
\mathbf{a} \otimes \mathbf{b}^\textsf{T}
\rightarrow
\begin{bmatrix}a_1 \\ a_2 \\ a_3 \\ a_4\end{bmatrix}  
\begin{bmatrix}b_1 & b_2 & b_3\end{bmatrix} = 
\begin{bmatrix}a_1b_1 & a_1b_2 & a_1b_3 \\ a_2b_1 & a_2b_2 & a_2b_3 \\ a_3b_1 & a_3b_2 & a_3b_3 \\ a_4b_1 & a_4b_2 & a_4b_3\end{bmatrix}
$$

```julia
using LinearAlgebra
a = [1, 2, 3]
v1 = [1, 2, 3]
v2 = [1, 2, 3]
outer_v = a * a'
cross_v = cross(v1, v2)

```

Output:
```
3-element Vector{Int64}:
 0
 0
 0
```

```julia
@assert outer_v == [1 2 3
                    2 4 6
                    3 6 9]
```

<!-- grade_id: cell-e6b6970ffe104df5, points: 1, grade: true, solution: false -->
```julia
@assert cross_v == [0, 0, 0]
```

```julia

```