# Основы линейной алгебры в Julia|
Author: Andreas Noack Jensen (MIT & JuliaComputing) (https://twitter.com/anoackjensen?lang=en)

Для начала, давайте объявим матрицу случайных величин

```julia
A = rand(1:4,3,3)
```

Output:
```
3×3 Matrix{Int64}:
 4  1  4
 1  3  1
 1  2  3
```

Объявим вектор единиц

```julia
x = fill(1.0, (3,)) # = fill(1.0, 3)
```

Output:
```
3-element Vector{Float64}:
 1.0
 1.0
 1.0
```

Обратите внимание, что $A$ имеет тип `Matrix{Int64}`, в то время как $x$ имеет тип `Vector{Float64}`. В Julia определены псевдонимы для `Vector{Type}=Array{Type,1}` и `Matrix{Type}=Array{Type,2}`. 

Многие базовые операторы выглядят так же, как в других языках програмирования.
#### Умножение

```julia
b = A*x
```

Output:
```
3-element Vector{Float64}:
 9.0
 5.0
 6.0
```

#### Транспорование
как и в *Matlab* `A'` возвращает Эрмитово-сопряженную матрицу (матрицу сопряженно-транспонированную)

```julia
A'
```

Output:
```
3×3 adjoint(::Matrix{Int64}) with eltype Int64:
 4  1  1
 1  3  2
 4  1  3
```

```julia
A
length(A[1,1:end])
```

Output:
```
3
```

```julia
x_5 = 5
x_5 = x+1
complex(1,3)
```

Output:
```
1 + 3im
```

Обычное транспонирование осуществляется при помощи

```julia
transpose(A)
```

Output:
```
3×3 transpose(::Matrix{Int64}) with eltype Int64:
 4  1  1
 1  3  2
 4  1  3
```

#### Транспонированное умножение
*Julia* позволяет записать без *

```julia
A'A
```

Output:
```
3×3 Matrix{Int64}:
 18   9  20
  9  14  13
 20  13  26
```

#### решение СЛАУ
Уравнение $Ax=b$ для ***квадратной*** $A$ решается при помощи оператора \.

```julia
b' /A
```

Output:
```
1×3 adjoint(::Vector{Float64}) with eltype Float64:
 2.13636  1.95455  -1.5
```

```julia

```

Также `A\b` возвращает *оптимальное в среднеквадратичном смысле* если у нас переопределенная система (т.е. матрица `A` "высокая")

```julia
Atall = rand(3, 2)
```

Output:
```
3×2 Matrix{Float64}:
 0.980481  0.229561
 0.980981  0.279894
 0.57909   0.341008
```

```julia
Atall\b
```

Output:
```
2-element Vector{Float64}:
 5.6743016502885455
 6.290770007511859
```

Если же система недоопределена, то оператор `\` возвращает  *решение с минимальной евклидовой нормой*

```julia
bshort = rand(2)
Ashort = rand(2, 3)
```

Output:
```
2×3 Matrix{Float64}:
 0.681196  0.959448  0.722411
 0.181582  0.235487  0.449536
```

```julia
Ashort\bshort
```

Output:
```
3-element Vector{Float64}:
 -0.3751175411602941
 -0.7614217392855187
  2.5574152678943722
```

# Модуль LinearAlgebra

Хоть и многие инструменты линейной алгебры доступны в Julia по умолчанию, существет стандартная библиотека `LinearAlgebra` которая предоставляет еще больше довольно важных функций. В частности, в ней имеются инструменты факторизации и некоторые типы для матриц определенного типа.

### Упражнения

#### 10.1 
Вычислите скалярное произведение вектора `v` с самиси собой и запишите результат в переменную `dot_v`.



```julia
v = [1,2,3]
```

```julia

```

<!-- grade_id: cell-913fef9b0d19cd52, points: 1, grade: true, solution: false -->
```julia
@assert dot_v == 14
```

#### 10.2 
Вычислите тензорное произведение вектора с самим собой и запишите полученную матрицу в переменную `outer_v`
$$
\mathbf{a} \otimes \mathbf{b}^\textsf{T}
\rightarrow
\begin{bmatrix}a_1 \\ a_2 \\ a_3 \\ a_4\end{bmatrix}  
\begin{bmatrix}b_1 & b_2 & b_3\end{bmatrix} = 
\begin{bmatrix}a_1b_1 & a_1b_2 & a_1b_3 \\ a_2b_1 & a_2b_2 & a_2b_3 \\ a_3b_1 & a_3b_2 & a_3b_3 \\ a_4b_1 & a_4b_2 & a_4b_3\end{bmatrix}
$$

```julia

```

```julia
@assert outer_v == [1 2 3
                    2 4 6
                    3 6 9]
```

<!-- grade_id: cell-e6b6970ffe104df5, points: 1, grade: true, solution: false -->
```julia
@assert cross_v == [0, 0, 0]
```

```julia

```