# Циклы

Содержание:
1. Цикл `while`
2. Цикл `for`
<br>

## Цикл while 

Синтаксис цикла `while` выглядит следующим образзом

```julia
while *условие*
    *тело цикла*
end
```

Например, мы можем использовать цикл `while`  что бы сосчитать элементы или пройтись по массиву.

```julia
n = 0
while n < 10
    n += 1
    println(n)
end
n
```

```julia
myfriends = ["Ted", "Robyn", "Barney", "Lily", "Marshall"]

i = 1
while i <= length(myfriends)
    friend = myfriends[i]
    println("Hi $friend, it's great to see you!")
    i += 1
end
```

## Цикла for

Синтаксис цикла `for` выглядит следующим образом

```julia
for *переменная* in *итерируемый объект*
    *тело цикла*
end
```

Назначения цикла `for` - перемещатся по массивам. Для привычнойго использования в духе языка `Pascal` можно использовать специальные массиво-подобные сущности - генераторы:

```julia
for n in 1:10
    println(n)
end
```

```julia
myfriends = ["Ted", "Robyn", "Barney", "Lily", "Marshall"]

for friend in myfriends
    println("Hi $friend, it's great to see you!")
end
```

Циклы могут быть вложенными, например, для обхода матрицы:

```julia
m, n = 5, 5
A = Array{Int32}(undef,m,n)
```

```julia
for j in 1:n
    for i in 1:m
        A[i, j] = i + j
    end
end
A
```

Также **Julia** позволяет такие записи вложенных циклов `for`

```julia
B = Array{Int32}(undef,m,n)
```

```julia
for j in 1:n, i in 1:m
    B[i, j] = i + j
end
B
```

Однако более в духе **Julia** будет использовать *генераторы массивов*

```julia
C = [i + j for i in 1:m, j in 1:n]
```

### Упражнения

#### 4.1 
Выведите квадраты чисел от 1 до 100.

```julia

```

#### 4.2 
Создайте словарь `squares`, в котором бы хранились числа и их квадраты

```julia

```

<!-- grade_id: cell-54834a2e4e122b70, points: 1, grade: true, solution: false -->
```julia
@assert squares[10] == 100
@assert squares[11] == 121
```

#### 4.3 
Создайте при помощи генератора массив `squares_arr`, содержащий квадраты числе от 1 до 100.

```julia

```

<!-- grade_id: cell-17d4631c51b8b4f7, points: 1, grade: true, solution: false -->
```julia
@assert length(squares_arr) == 100
@assert sum(squares_arr) == 338350
```

#### 4.4

Создайте массив `primes`, содержащий все простые числа от 1 до 100. (Используйте цикл for)

```julia
primes = Array{Int16,1}(undef, 0)
for i in 2:100
    b = false;
    for num in primes
        if i%num == 0
            b = true
            break
        end
    end
    if !b
        push!(primes,i);
    end
end
primes
```

Output:
```
25-element Vector{Int16}:
  2
  3
  5
  7
 11
 13
 17
 19
 23
 29
  ⋮
 59
 61
 67
 71
 73
 79
 83
 89
 97
```