# Циклы

Содержание:
1. Цикл `while`
2. Цикл `for`
<br>

## Цикл while 

Синтаксис цикла `while` выглядит следующим образзом

```julia
while *условие*
    *тело цикла*
end
```

Например, мы можем использовать цикл `while`  что бы сосчитать элементы или пройтись по массиву.

```julia
n = 0
while n < 10
    n += 1
    println(n)
end
n
```

Output:
```
1
2
3
4
5
6
7
8
9
10
10
```

```julia
myfriends = ["Ted", "Robyn", "Barney", "Lily", "Marshall"]

i = 1
while i <= length(myfriends)
    friend = myfriends[i]
    println("Hi $friend, it's great to see you!")
    i += 1
end
```

Output:
```
Hi Ted, it's great to see you!
Hi Robyn, it's great to see you!
Hi Barney, it's great to see you!
Hi Lily, it's great to see you!
Hi Marshall, it's great to see you!
```

## Цикла for

Синтаксис цикла `for` выглядит следующим образом

```julia
for *переменная* in *итерируемый объект*
    *тело цикла*
end
```

Назначения цикла `for` - перемещатся по массивам. Для привычнойго использования в духе языка `Pascal` можно использовать специальные массиво-подобные сущности - генераторы:

```julia
for n in 1:length(myfriends)
    println(n)
end
```

Output:
```
1
2
3
4
5
```

```julia
myfriends = ["Ted", "Robyn", "Barney", "Lily", "Marshall"]

for friend in myfriends
    println("Hi $friend, it's great to see you!")
end
```

Output:
```
Hi Ted, it's great to see you!
Hi Robyn, it's great to see you!
Hi Barney, it's great to see you!
Hi Lily, it's great to see you!
Hi Marshall, it's great to see you!
```

Циклы могут быть вложенными, например, для обхода матрицы:

```julia
m, n = 5, 5
A = Array{Int32}(undef,m,n)
```

Output:
```
5×5 Matrix{Int32}:
  -867027056        32764  -1418856432        32763  -1430974272
         450  -1396919056          450  -1415902320          450
 -1396919168          450   -255373920          450   -184834864
         450  -1396919056        32763   -255374160        32763
   214542528          450   -255373920        32763            0
```

```julia
for j in 1:n
    for i in 1:m
        A[i, j] = i + j
    end
end
A
```

Output:
```
5×5 Matrix{Int32}:
 2  3  4  5   6
 3  4  5  6   7
 4  5  6  7   8
 5  6  7  8   9
 6  7  8  9  10
```

Также **Julia** позволяет такие записи вложенных циклов `for`

```julia
B = Array{Int32}(undef,m,n)
```

Output:
```
5×5 Matrix{Int32}:
 -1411607152          450  -1411606832          450  -1411606512
         450  -1411606960          450  -1411606640          450
 -1411607088          450  -1411606768          450  -1411606448
         450  -1411606896          450  -1411606576          450
 -1411607024          450  -1411606704          450     16842752
```

```julia
for j in 1:n, i in 1:m
    B[i, j] = i + j
end
B
```

Output:
```
5×5 Matrix{Int32}:
 2  3  4  5   6
 3  4  5  6   7
 4  5  6  7   8
 5  6  7  8   9
 6  7  8  9  10
```

Однако более в духе **Julia** будет использовать *генераторы массивов*

```julia
C = [i + j for i in 1:m, j in 1:n]
```

Output:
```
5×5 Matrix{Int64}:
 2  3  4  5   6
 3  4  5  6   7
 4  5  6  7   8
 5  6  7  8   9
 6  7  8  9  10
```

### Упражнения

#### 4.1 
Выведите квадраты чисел от 1 до 100.

```julia
for i in 1:100
    println(i^2)
end
```

Output:
```
1
4
9
16
25
36
49
64
81
100
121
144
169
196
225
256
289
324
361
400
441
484
529
576
625
676
729
784
841
900
961
1024
1089
1156
1225
1296
1369
1444
1521
1600
1681
1764
1849
1936
2025
2116
2209
2304
2401
2500
2601
2704
2809
2916
3025
3136
3249
3364
3481
3600
3721
3844
3969
4096
4225
4356
4489
4624
4761
4900
5041
5184
5329
5476
5625
5776
5929
6084
6241
6400
6561
6724
6889
7056
7225
7396
7569
7744
7921
8100
8281
8464
8649
8836
9025
9216
9409
9604
9801
10000
```

#### 4.2 
Создайте словарь `squares`, в котором бы хранились числа и их квадраты

```julia
squares = Dict()
for i in 1:100
    squares[i] = i^2
end
squares
```

Output:
```
Dict{Any, Any} with 100 entries:
  5  => 25
  56 => 3136
  35 => 1225
  55 => 3025
  60 => 3600
  30 => 900
  32 => 1024
  6  => 36
  67 => 4489
  45 => 2025
  73 => 5329
  64 => 4096
  90 => 8100
  4  => 16
  13 => 169
  54 => 2916
  63 => 3969
  86 => 7396
  91 => 8281
  ⋮  => ⋮
```

<!-- grade_id: cell-54834a2e4e122b70, points: 1, grade: true, solution: false -->
```julia
@assert squares[10] == 100
@assert squares[11] == 121
```

#### 4.3 
Создайте при помощи генератора массив `squares_arr`, содержащий квадраты числе от 1 до 100.

```julia
squares_arr = [i^2 for i in 1:100]
```

Output:
```
100-element Vector{Int64}:
     1
     4
     9
    16
    25
    36
    49
    64
    81
   100
     ⋮
  8464
  8649
  8836
  9025
  9216
  9409
  9604
  9801
 10000
```

<!-- grade_id: cell-17d4631c51b8b4f7, points: 1, grade: true, solution: false -->
```julia
@assert length(squares_arr) == 100
@assert sum(squares_arr) == 338350
```

#### 4.4

Создайте массив `primes`, содержащий все простые числа от 1 до 100. (Используйте цикл for)

```julia
primes = Array{Int16,1}(undef,0)
for i in 2:100
    b = false
    for num in primes
        if i % num == 0
            b = true
            break
        end
    end
    if !b
        push!(primes, i)
    end
end
primes


```

Output:
```
25-element Vector{Int16}:
  2
  3
  5
  7
 11
 13
 17
 19
 23
 29
  ⋮
 59
 61
 67
 71
 73
 79
 83
 89
 97
```